// variables2.rs
// Make me compile! Execute the command `rustlings hint variables2` if you want a hint :)

<<<<<<< HEAD
// I AM NOT DONE

fn main() {
    let x;
=======

fn main() {
    let x = 10;
>>>>>>> 90fab5c (2022/7/17/14.12)
    if x == 10 {
        println!("Ten!");
    } else {
        println!("Not ten!");
    }
}
