// variables5.rs
<<<<<<< HEAD
// Make me compile! Execute the command `rustlings hint variables5` if you want a hint :)

// I AM NOT DONE
=======
// Make me compile! Execute the command `rustlings hint variables5` if you want a hint :rustlings hint variables5)

>>>>>>> 90fab5c (2022/7/17/14.12)

fn main() {
    let number = "T-H-R-E-E"; // don't change this line
    println!("Spell a Number : {}", number);
<<<<<<< HEAD
    number = 3;
=======
    let number = 3;
>>>>>>> 90fab5c (2022/7/17/14.12)
    println!("Number plus two is : {}", number + 2);
}
