// variables3.rs
// Make me compile! Execute the command `rustlings hint variables3` if you want a hint :)

<<<<<<< HEAD
// I AM NOT DONE

fn main() {
    let x = 3;
=======

fn main() {
    let mut x = 3;
>>>>>>> 90fab5c (2022/7/17/14.12)
    println!("Number {}", x);
    x = 5; // don't change this line
    println!("Number {}", x);
}
