// primitive_types3.rs
// Create an array with at least 100 elements in it where the ??? is.
// Execute `rustlings hint primitive_types3` for hints!

<<<<<<< HEAD
// I AM NOT DONE

fn main() {
    let a = ???
=======

fn main() {
    let a = [50,100];
>>>>>>> 90fab5c (2022/7/17/14.12)

    if a.len() >= 100 {
        println!("Wow, that's a big array!");
    } else {
        println!("Meh, I eat arrays like that for breakfast.");
    }
}
