// primitive_types1.rs
// Fill in the rest of the line that has code missing!
// No hints, there's no tricks, just get used to typing these :)

<<<<<<< HEAD
// I AM NOT DONE
=======
>>>>>>> 90fab5c (2022/7/17/14.12)

fn main() {
    // Booleans (`bool`)

    let is_morning = true;
    if is_morning {
        println!("Good morning!");
    }

<<<<<<< HEAD
    let // Finish the rest of this line like the example! Or make it be false!
=======
    let is_evening = true;// Finish the rest of this line like the example! Or make it be false!
>>>>>>> 90fab5c (2022/7/17/14.12)
    if is_evening {
        println!("Good evening!");
    }
}
