// tests2.rs
// This test has a problem with it -- make the test compile! Make the test
// pass! Make the test fail! Execute `rustlings hint tests2` for hints :)

<<<<<<< HEAD
// I AM NOT DONE
=======
>>>>>>> 90fab5c (2022/7/17/14.12)

#[cfg(test)]
mod tests {
    #[test]
    fn you_can_assert_eq() {
<<<<<<< HEAD
        assert_eq!();
=======
        assert_eq!(1,1);
>>>>>>> 90fab5c (2022/7/17/14.12)
    }
}
