// modules3.rs
// You can use the 'use' keyword to bring module paths from modules from anywhere
// and especially from the Rust standard library into your scope.
// Bring SystemTime and UNIX_EPOCH
// from the std::time module. Bonus style points if you can do it with one line!
// Make me compile! Execute `rustlings hint modules3` for hints :)

<<<<<<< HEAD
// I AM NOT DONE

// TODO: Complete this use statement
use ???
=======

// TODO: Complete this use statement
use std::time::UNIX_EPOCH;
use std::time::SystemTime;
>>>>>>> 90fab5c (2022/7/17/14.12)

fn main() {
    match SystemTime::now().duration_since(UNIX_EPOCH) {
        Ok(n) => println!("1970-01-01 00:00:00 UTC was {} seconds ago!", n.as_secs()),
        Err(_) => panic!("SystemTime before UNIX EPOCH!"),
    }
}
