// macros2.rs
// Make me compile! Execute `rustlings hint macros2` for hints :)

<<<<<<< HEAD
// I AM NOT DONE

fn main() {
    my_macro!();
}

=======
#[warn(unused_macros)]
>>>>>>> 90fab5c (2022/7/17/14.12)
macro_rules! my_macro {
    () => {
        println!("Check out my macro!");
    };
}
<<<<<<< HEAD
=======
fn main() {
    my_macro!();
}
>>>>>>> 90fab5c (2022/7/17/14.12)
