// macros1.rs
// Make me compile! Execute `rustlings hint macros1` for hints :)

<<<<<<< HEAD
// I AM NOT DONE

=======
#[warn(unused_macros)]
>>>>>>> 90fab5c (2022/7/17/14.12)
macro_rules! my_macro {
    () => {
        println!("Check out my macro!");
    };
}

fn main() {
<<<<<<< HEAD
    my_macro();
=======
    my_macro!();
>>>>>>> 90fab5c (2022/7/17/14.12)
}
