// enums1.rs
// Make me compile! Execute `rustlings hint enums1` for hints!

<<<<<<< HEAD
// I AM NOT DONE
=======
>>>>>>> 90fab5c (2022/7/17/14.12)

#[derive(Debug)]
enum Message {
    // TODO: define a few types of messages as used below
<<<<<<< HEAD
=======
    Quit,
    Echo,
    Move,
    ChangeColor,
>>>>>>> 90fab5c (2022/7/17/14.12)
}

fn main() {
    println!("{:?}", Message::Quit);
    println!("{:?}", Message::Echo);
    println!("{:?}", Message::Move);
    println!("{:?}", Message::ChangeColor);
}
