// functions5.rs
// Make me compile! Execute `rustlings hint functions5` for hints :)

<<<<<<< HEAD
// I AM NOT DONE
=======
>>>>>>> 90fab5c (2022/7/17/14.12)

fn main() {
    let answer = square(3);
    println!("The answer is {}", answer);
}

fn square(num: i32) -> i32 {
<<<<<<< HEAD
    num * num;
=======
    num * num
>>>>>>> 90fab5c (2022/7/17/14.12)
}
