// functions3.rs
// Make me compile! Execute `rustlings hint functions3` for hints :)

<<<<<<< HEAD
// I AM NOT DONE

fn main() {
    call_me();
=======

fn main() {
    call_me(4);
>>>>>>> 90fab5c (2022/7/17/14.12)
}

fn call_me(num: u32) {
    for i in 0..num {
        println!("Ring! Call number {}", i + 1);
    }
}
