// functions2.rs
// Make me compile! Execute `rustlings hint functions2` for hints :)

<<<<<<< HEAD
// I AM NOT DONE
=======
>>>>>>> 90fab5c (2022/7/17/14.12)

fn main() {
    call_me(3);
}

<<<<<<< HEAD
fn call_me(num:) {
=======
fn call_me(num:i32) {
>>>>>>> 90fab5c (2022/7/17/14.12)
    for i in 0..num {
        println!("Ring! Call number {}", i + 1);
    }
}
